# COGS 108 - Assignment 2: Data Exploration

# Important

- Rename this file to 'A2_$####.ipynb' (filled in with your unique course ID) before you submit it. Submit it to TritonED.
- This assignment has hidden tests: tests that are not visible here, but that will be run on your submitted assignment.
    - This means passing all the tests you can see in the notebook here does not guarantee you have the right answer!
- Each coding question in this assignment only requires a small amount of code, about 1-3 lines. 
    - If you find yourself writing much more than that, you might want to reconsider your approach. 
    - A reminder that the Tutorials notebooks often contain similar examples to those used in the assignments.
- This assignment also has some written answers, in which you will write a short response (in text, not code). 
    - These questions are referred to as 'W#', and will be followed by a markdown cell in which you can write your answer. 
    - These are short answer questions, prompting you to think about the approaches we are using. 
    - You answers can and should be fairly brief, typically about 1-2 sentences.

```python
# Imports
#  Note: these are all the imports you need! Do not import any other functions / packages
%matplotlib inline

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

## Part 1 - Data Wrangling

For this assignment, you are provided with two data files:
- age_steps.csv
- name_income_id.json

<!-- grade_id: load_steps, points: 0, grade: false, solution: true -->
```python
# 1a) Load the csv file into a pandas dataframe. Call it 'df_steps'.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_load_steps, points: 0.5, grade: true, solution: false -->
```python
# Tests for 1a

assert isinstance(df_steps, pd.DataFrame)

```

<!-- grade_id: load_income, points: 0, grade: false, solution: true -->
```python
# 1b) Load the json file into a (different) pandas dataframe. Call it 'df_income'.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_load_income, points: 0.5, grade: true, solution: false -->
```python
# Tests for 1b

assert isinstance(df_income, pd.DataFrame)

```

<!-- grade_id: drop_extra_cols, points: 0, grade: false, solution: true -->
```python
# 1c) Drop the 'first_name' and 'last_name' columns from the df_income data frame.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_drop_ext_cols, points: 0.25, grade: true, solution: false -->
```python
# Tests for 1c

assert 'first_name' not in df_income.columns
assert 'last_name' not in df_income.columns
```

<!-- grade_id: merge_dat, points: 0, grade: false, solution: true -->
```python
# 1d) Combine the two dataframes into one combined, new dataframe. Call it 'df'.
#  Both dataframes should have a column called 'id' use this column to align the datasets (merge on 'id')
#  Note: For this part, it doesn't in which order you merge the dataframes.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_merge_dat, points: 0.5, grade: true, solution: false -->
```python
# Tests for 1d

assert isinstance(df, pd.DataFrame)
assert set(df.columns) == set(['id', 'income', 'age', 'steps'])
```

<!-- grade_id: reorder_cols, points: 0, grade: false, solution: true -->
```python
# 1e) Re-Order the Column ID's to be [id, age, steps, income]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_reorder_cols, points: 0.25, grade: true, solution: false -->
```python
# Tests for 1e

assert list(df.columns) == ['id', 'age', 'steps', 'income']
```

<!-- grade_id: cell-21fcd98a93bb7fa1, points: 0, grade: false, solution: true -->
```python
# 1f) Make two different copies of 'df'. 
# Call the first one 'df2', and the second one 'df3'.
# Note: Simply writing 'df2 = df' and 'df3 = df' will NOT work!
#   A quick google search reveals how to properly copy dataframes.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: import, points: 0.5, grade: true, solution: false -->
```python
# Tests for 1f

assert(df2 is not df)
assert(df3 is not df)
```

## Part 2 - Data Cleaning

Before analysis, data must be cleaned of missing or unusable data. 

Now that the datasets are joined into a single dataframe, the next steps are to remove or replace any bad data before moving onto any analyses. 

<!-- grade_id: check_n_nan, points: 0, grade: false, solution: true -->
```python
# 2a) First check how many missing data (NaNs) there are in the income data.
#  Save this number into a variable called 'n_nan'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_check_n_nan, points: 0.25, grade: true, solution: false -->
```python
# Tests for 2a

assert(n_nan)

```

<!-- grade_id: rm_nans, points: 0, grade: false, solution: true -->
```python
# 2b) Remove all rows from df that have missing data. 
#  In other words, remove all rows with NaN values.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_nans, points: 0.5, grade: true, solution: false -->
```python
# Tests for 2b

assert sum(np.isnan(df['income'])) == 0
```

<!-- grade_id: cell-07562783e1097577, points: 0, grade: false, solution: true -->
```python
# 2c) Create a variable called `mean_income` which holds the average income in df

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0087330817bb378a, points: 0.25, grade: true, solution: false -->
```python
# Tests for 2c

assert mean_income

```

<!-- grade_id: cell-ae9e4583c7334cef, points: 0, grade: false, solution: true -->
```python
# 2d) In df2, replace all missing values with `mean_income`

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e6b939ab097a08ac, points: 0.5, grade: true, solution: false -->
```python
# Tests for 2d

assert not df2['income'].hasnans

```

### W1 - Missing Data

You just used two different approaches for dealing with missing data: dropping the rows and replacing with the mean. 

For each approach, briefly describe below when you think it would be appropriate to use this approach. 

When is best to use one approach over the other?

<!-- grade_id: cell-50295605e63eb00b, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 3: Data Visualization

Here we will start exploring our data visually, using some plots. 

In particular, we will plot the distributions of each of our variables, in such a way as to let us get a sense of the data, and look for any problems. 

Pandas comes with some plotting capabilities. We will also be making histograms, which can be done with matplotlib.

Here's an example of making histograms: http://matplotlib.org/1.2.1/examples/pylab_examples/histogram_demo.html

This example has more than you need: 
- Titles, axis labels, etc are optional (but recommended)
- You do no need to add a 'best fit' line.

Notes:

- Everywhere that we ask you to create a plot, make sure to leave the 'plt.gcf()' line at the end of the cell. 
    - This is used to get a figure handle in order to programmatically evaluate your plots.
- For all your histograms, use 25 bins

<!-- grade_id: hist_age, points: 0, grade: false, solution: true -->
```python
# 3a) Plot a histogram of the age variable

# YOUR CODE HERE
raise NotImplementedError()

f1 = plt.gcf()
```

<!-- grade_id: test_hist_age, points: 0.25, grade: true, solution: false -->
```python
# Tests for 3a

assert f1.gca().has_data()
```

<!-- grade_id: hist_steps, points: 0, grade: false, solution: true -->
```python
# 3b) Plot a histogram of the steps variable

# YOUR CODE HERE
raise NotImplementedError()

f2 = plt.gcf()
```

<!-- grade_id: test_hist_steps, points: 0.25, grade: true, solution: false -->
```python
# Tests for 3b

assert f2.gca().has_data()
```

<!-- grade_id: cell-3b187fca71faa424, points: 0, grade: false, solution: true -->
```python
# 3c) Plot a histogram of the income variable

# YOUR CODE HERE
raise NotImplementedError()

f3 = plt.gcf()
```

<!-- grade_id: cell-f9b252fbb7a92d40, points: 0.25, grade: true, solution: false -->
```python
# Tests for 3c

assert f2.gca().has_data()
```

<!-- grade_id: cell-95bdb279f6c0f560, points: 0, grade: false, solution: true -->
```python
# 3d) Plot the data using pandas scatter_matrix, saving it to output variable name 'fd'.
#  Only plot the 'age', 'steps', and 'income' variables.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c5520747d92f0e26, points: 0.25, grade: true, solution: false -->
```python
# Tests for 3d

assert len(fd[0]) == 3
assert np.all(ax.has_data() for ax in fd[0])
```

### W2 - What do you notice about data now that you've visualized it?

For each of 'steps', 'income' and 'age', write a quick comment on each of anything that you notice from the above visualizations, that may need to be addressed before moving onto data analysis. 

In particular:
- Are there outliers or bad values we need to fix? 
- Might we want to transform any of the distributions?

<!-- grade_id: notice_steps, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 4: Data Pre-Processing

In the above sections, we did some basic data cleaning, and visualization. 

In practice, these two components of an analysis pipeline are often combined into an iterative approach, going back and forth between looking at the data, checking for issues, and then cleaning the data. 

In this next section we will continue with an iterative procedure of data cleaning and visualization, addressing some issues that we can start to see when visualizing the data.

<!-- grade_id: check_n_neg, points: 0, grade: false, solution: true -->
```python
# 4a) Dealing with an impossible value in 'steps'. 
#  Check how many -1 values there are in steps. 
#  Save the number into a variable called 'n_neg'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_check_n_neg, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4a

assert(n_neg)

```

<!-- grade_id: rm_neg, points: 0, grade: false, solution: true -->
```python
# 4b) Remove all rows where steps is equal to -1

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_neg, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4b

assert sum(df['steps'] == -1) == 0
```

### W3 - Transforming Distributions

Once thing you might have noticed is that the income is not normally distributed. 

In the next steps, we will address this by log transforming the income data. First though, we will have to deal with any income values that are 0. Note that these values are not impossible values, but may, for example, reflect people who are unemployed. 

Why are are we log-transforming the data, and how does that help us? Why do we have to do something special with 0's?

Hint: calculate log10 of zero, and see what the answer is. 

<!-- grade_id: cell-132cc0868f89e090, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: recoded_zero, points: 0, grade: false, solution: true -->
```python
# 4c) For all rows in which income is equal to zero, recode income to be equal to 1

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_recoded_zero, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4c

assert sum(df['income'] == 0.) == 0
```

<!-- grade_id: log, points: 0, grade: false, solution: true -->
```python
# 4d) Transform income using a log10 transform

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_log, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4d

assert np.isclose(min(df['income']), 0.0, 0.05)
assert np.isclose(max(df['income']), 5.62, 0.05)
```

<!-- grade_id: replot_log, points: 0, grade: false, solution: true -->
```python
# 4e) Replot the histogram for income data after the data transformation

# YOUR CODE HERE
raise NotImplementedError()

f4 = plt.gcf()
```

<!-- grade_id: test_replot_hist, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4e

assert f4.gca().has_data()
```

<!-- grade_id: rm_young, points: 0, grade: false, solution: true -->
```python
# 4f) We might also have certain regulations or restrictions that we need to follow about the data. 
# Here, we will only analyze adults. Remove all rows where age is less than 18.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_young, points: 0.25, grade: true, solution: false -->
```python
# Tests for 4f

assert min(df['age']) >= 18
```

## Part 5 Outliers

So far, we have dealt with the most clearly wrong or impossible data values. 

We might still might need to deal with more subtle outliers: values that seem abnormally high or low. These values may seem suspicious, and they may skew our analyses. 

In this step, we are going to deal with outlier values by removing all rows in which the data for the steps column exceeds +/- 3 standard deviations above or below the mean.

Note that when you perform this kind of outlier removal can have large impacts. For example, the procedure we take here might give quite different results if we did it before removing the -1 values in steps. 

<!-- grade_id: rm_outliers_setup, points: 0, grade: false, solution: true -->
```python
# 5a) For the steps data: 
# - Calculate what three times the standard deviation is, and save it to variables called 'std3'. 
#    Note: you must use pandas to calculate the standard devation!
#    Pandas & Numpy use slightly different conventions for std, so using numpy here will get a different value
# - Using 'std3', create new variables, called std_max' and 'std_min'.
#    These will store our cut-off values: 3 std above and below the mean respectively. 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_outliers_setup, points: 0.5, grade: true, solution: false -->
```python
# Tests for 5a

assert(std3)
assert(std_min)
assert(std_max)

```

<!-- grade_id: rm_outliers, points: 0, grade: false, solution: true -->
```python
# 5b) Now use these values to remove any rows with steps data above or below these thresholds

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tests_rm_outliers, points: 0.5, grade: true, solution: false -->
```python
# Tests for 5b

assert np.all(df['steps'])

```

```python
# CHECK IN
# Check that the following test passes - if it fails, you have the wrong amount of data left. 
#  If so, go back and try to fix your cleaning steps, else your analysis answers will be wrong.
assert df.shape == (1463, 4)
```

## Part 6 - Basic Analyses

Now that we have wrangled and cleaned our data, we can start doing some simple analyses. 

Here we will explore some basic descriptive summaries of our data, look into the inter-relations (correlations) between variables, and ask some simple questions about potentially interesting subsets of our data. 

<!-- grade_id: cell-ae86332aaa14d006, points: 0, grade: false, solution: true -->
```python
# 6a) Descriptive summary
# Use the 'describe' pandas method to check a descriptive summary of the data
#   Save the dataframe generated by describe to a new variable called 'desc'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ff715fe5f2413f68, points: 0.5, grade: true, solution: false -->
```python
# Tests for 6a

assert isinstance(desc, pd.DataFrame)
```

<!-- grade_id: calc_corrs, points: 0, grade: false, solution: true -->
```python
# 6b) Calculate the pairwise correlations between all variables. 
#  Notes: do this with a pandas method. Keep all columns (including ID).
#  Assign the result (which should be a dataframe) to a variable called 'corrs'.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_calc_corrs, points: 0.25, grade: true, solution: false -->
```python
# Tests for 6b

assert isinstance(corrs, pd.DataFrame)
assert corrs.shape == (4, 4)
```

<!-- grade_id: cell-0beee51a9a82e017, points: 0, grade: false, solution: true -->
```python
# 6c) Checking correlations
# Answer the following questions by setting to answer variable to 'age', 'steps', or 'income'
# Which variable is most correlated with age? Set the answer in a variable called 'age_corr'. 
# Which variable is most correlated with income? Set the answer in a variable called 'inc_corr'. 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d720528b12bb00be, points: 0.25, grade: true, solution: false -->
```python
# Tests for 6c

assert age_corr
assert inc_corr

```

<!-- grade_id: cell-0cd6f99f2dbdb767, points: 0, grade: false, solution: true -->
```python
# 6d) How many steps would you have to walk to be in the top 10% of walkers
#  Save you answer as a variable called 'top_walker'
#  Hint: check out the 'quantile' method.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cfed40af74908146, points: 0.25, grade: true, solution: false -->
```python
# Tests for 6d

assert top_walker

```

<!-- grade_id: cell-02b8f7a3ac9705f0, points: 0, grade: false, solution: true -->
```python
# 6e) What is the average income for people over the age of 45?
# Save your response in a variable called 'old_income'.
# Note: for this response, keep the answer as the logged value (do no 'un-log' the answer).

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-332be768960bcc82, points: 0.25, grade: true, solution: false -->
```python
# Tests for 6e
assert old_income

```

# Part 7 - Predictions

A frequent goal of data analysis is to understand so that we can make predictions about future or unseen data points. 

Here we will explore some basic predictions, looking into whether we might be able to predict income from our other variables.

Notes:
- Use the polyfit function from numpy, in the way that we did in 'SectionMaterials/02-DataAnalysis
- The predictions for income should be in dollars (so you'll have to 'un-log' the prediction)

<!-- grade_id: model_1, points: 0, grade: false, solution: true -->
```python
# 7a) Use polyfit to fit a 1-degree linear model, predicting income from age.
#  Call the output parameters 'a1' and 'b1'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_model_1, points: 0.5, grade: true, solution: false -->
```python
# Tests for 7a

assert(a1)
assert(b1)

```

<!-- grade_id: pred_1, points: 0, grade: false, solution: true -->
```python
# 7b) Use the model parameters from 4a to predict the income of a 75 year old. 
#  Call your prediction 'pred1'.
#  Note: Convert the prediction to be in real dollars (not log dollars)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_pred_1, points: 0.5, grade: true, solution: false -->
```python
# Tests for 7b

assert(pred1)

```

<!-- grade_id: model_2, points: 0, grade: false, solution: true -->
```python
# 7c) Use polyfit to fit a 1-degree linear model, predicting income from steps. 
#  Call the output parameters 'a2' and 'b2'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_model_2, points: 0.25, grade: true, solution: false -->
```python
# Tests for 7c

assert(a2)
assert(b2)

```

<!-- grade_id: pred_2, points: 0, grade: false, solution: true -->
```python
# 7d) Predict the income of someone who took 20,000 steps. Call your prediction 'pred2'.
#  Note: Convert the prediction to be in real dollars (not log dollars)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_pred_2, points: 0.25, grade: true, solution: false -->
```python
# Test for 7d

assert(pred2)

```

<!-- grade_id: better_model, points: 0, grade: false, solution: true -->
```python
# 7e) Based on your analysis in Part 3, which model do you think is more likely to have made a good prediction?
#  Create a variable called 'better_model', and set it equal to 'age' or 'steps', whichever your answer is.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_better_model, points: 0.5, grade: true, solution: false -->
```python
assert better_model

```

### W4 - Why do you think (whichever model you answered) is better?

<!-- grade_id: model_ans, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## 8  - Revisiting Missing Values

Originally, we dealt with missing values by either dropping rows that contained them or by replacing them with the mean value of that column. 

Another approach one can take is to predict missing values, based on the rest of the available data.

<!-- grade_id: cell-ee32dbc3740c0f4c, points: 0, grade: false, solution: true -->
```python
# 8a) Create a function named `predict_income` that takes one parameter, `age`.
#
# This function will take in as input an individual's age, and return their predicted income.
# This function should use your parameters `a1` and `b1` to do so.
# Also, make sure that the function returns predictions in dollars (not logged values).
# Note that this function should require no more than 2 or 3 lines.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-89038da0487305f5, points: 0.5, grade: true, solution: false -->
```python
# Tests for 8a

assert callable(predict_income)

```

```python
# Using the function you defined, the following code will apply the function to the dataframe (this code provided)
#  Now, the variable 'my_predictions' holds the predicted income for every row in df3. 
my_predictions = df3['age'].apply(predict_income)
```

### W5: Pandas .apply()

The code that you see in the line above uses a pandas method called 'apply'. 

It is a very powerful method, and can be useful in many situations.

Look into the '.apply()' method (remember you can use '?' in notebooks to see documentation). 

In the cell below, briefly explain what pandas '.apply()' method is, and what it does.

<!-- grade_id: cell-e823042187301644, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-5bef3033a299977a, points: 0, grade: false, solution: true -->
```python
# 8b) Create a new column in df3 named `predicted_income`. 
#   This column should includes the income that your model (a1, b1) predicts for each individual given their age. 
# Hint: you can use 'my_predictions', from above. 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9fc4e69c2fa2ed33, points: 0.5, grade: true, solution: false -->
```python
# Tests for 8b

assert 'predicted_income' in df3.columns

```

### W6: When do you think that it is NOT appropiate to use prediction models to fill in missing data? 

<!-- grade_id: cell-d0b519c4be7a9ca9, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Done! Upload this notebook to TritonED