# COGS 108 - Assignment 1: Setting Up

This assignment covers Python, git, and Github. 

Visit the relevant tutorials on https://github.com/COGS108/Tutorials for guidance for these topics.

This assignment is out of 12 points, worth 12% of your grade.

## IMPORTANT - FILE NAMING

For this and all other assignments you will be renaming them using your unique 'course ID'. 

This 'course ID' is comprised of the first letter of your last name, followed by the last 4 numbers of your student ID. 

So, a hypothetical student, A. Turing, with student number A10010101 would have course ID: T0101

### Rename this file to 'A1_$####.ipynb'


## ALSO IMPORTANT

### In the markdown cell below, write in your Student ID, unique course ID, Github username. 

The details you enter below will be how we get your account details to check that you have completed the Github portions of this assignment.

Please do not remove anything from the cell below, just add in the requested information to each line.

My Student ID number is: 

My Course ID is:

^This is the same 'course ID' you created above, the same used in the name of this file.

My Github username is: 

^Make sure this is your username, not the email address you used to create the account.

The following couple cells are checks that you are running the right tools for these assignments.

If either of them fail, you will need to update some things. If you are not sure how, ask in office hours.

```python
# These assignments require Python 3 (not Python 2).
#  Run this cell to check your version of Python.
import sys
assert sys.version_info[0] == 3, "You must use Python 3 (preferably 3.6) for these assignments."
```

```python
# Run this cell to check your version of IPython. 
#  If you get an error (the message prints out), stop and upgrade Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Part 1 (6 points)

Follow along on the instructions in the assignment description to do Part 1 of the assignment with git/Github.

## Part 2 (6 points)

This part of the assignment is focused on some practice with Python, and with practicing working with the format of the assignments. 

This class assumes some prior knowledge of Python. In the following questions, you will need to work with basic (standard library) data types (floats, lists, dictionaries, etc.) and control flow (conditionals, loops, functions, etc). If the questions in this section are totally unfamiliar to you, you may need revisit some practice materials to catch up with some of the programming. 

Through these questions, we will also prompt you to use a couple slightly more advanced standard library functions (for example, 'enumerate' and 'zip'), that may be new to you.

Each question should be answerable with a relatively small number of lines of code, up to about 5-7 lines at most. 

If you are having any trouble, remember to visit the course tutorials (https://github.com/COGS108/Tutorials). Assignment questions often follow the structure of examples provided in the tutorials, and a large number of relevant links and external materials are also indexed in the tutorials. 

### How to complete assignments

Whenever you see:

YOUR CODE HERE <br>
raise NotImplementedError()

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

Recall that any cell with 'assert' statements in it is a test cell. You should not try to change or delete these cells.

Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

```python
# PRINTING VARIABLES
# A reminder that you can (and should) print and check variables as you go.
#  This allows you to check what values they hold, and debug if anything unexpected happens.

# Define a variable
math_result = 2 * 4

# Print out the value(s) of a variable.
print(math_result)
```

<!-- grade_id: cell-4db8c933a259ed1b, points: 0, grade: false, solution: true -->
```python
# Q1 - Defining variables (1 point)

# Define a list, called 'var_a', that contains individual letters a-j (inclusively).
# Define a tuple, called 'var_b', that contains the numbers 1-10 (inclusively).

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88582fbb27969b63, points: 1, grade: true, solution: false -->
```python
# Tests for Q1

# These tests check the variables are defined
assert var_a
assert var_b

# These tests check that the variables are the right data types
assert isinstance(var_a, list)
assert isinstance(var_b, tuple)

# These tests check the variables have the right values
assert var_a == ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
assert var_b == (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)
```

<!-- grade_id: cell-a0c9ee8ea65481d3, points: 0, grade: false, solution: true -->
```python
# Q2 - Defining Variables, Part II (1 point)

# Create a dictionary, and call it 'dictionary'.
# Fill the dictionary such that each key is an element from 'var_a',
#  and each value is the corresponding element from 'var_b'.
# Hint: one useful function for this you can use is 'zip'.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-611de86b26434b5e, points: 1, grade: true, solution: false -->
```python
# Tests for Q2
assert dictionary

# There are also some hidden tests that check the contents of dictionary
```

<!-- grade_id: cell-5c803671612c731a, points: 0, grade: false, solution: true -->
```python
# Q3 - Control Flow (1 point)

# These variables are provided to you.
lst = [1, 2, 3, 4, 5, 6]
inds = []

# Loop through the provided list 'lst'. For each element, check if it is an even number.
#  If the element is an even number, add the INDEX of that element to the list 'inds'.
#    Note that you are adding the index to 'inds', not the element itself.
# Hint: to check if a number is even, you can use the modulo '%' operator.
# Hint: to loop through an iterable, keeping track of the index, you can use the 'enumerate' function.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-90cd314531614439, points: 1, grade: true, solution: false -->
```python
# Test for Q3

assert inds

# There is also a hidden test that checks the contents of 'inds'
```

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# Q4 - Functions (1 point)

# Write a function that takes two inputs, adds them together, and returns the result.
#  Call the function 'add2'.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: add2-tests, points: 1, grade: true, solution: false -->
```python
# Tests for Q4

assert add2(1, 2) == 3
assert add2(-1, 1) == 0
```

<!-- grade_id: cell-50b7166b6554cfff, points: 0, grade: false, solution: true -->
```python
# Q5 - Putting it all together (1 point)

# This variable is provided to you.
other_dictionary = dict()

# We will now fill 'other_dictionary', from data.
#  Loop through the key-value pairs in 'dictionary'.
#    For each key-value pair, if the value is odd, use the 'add2' to add 1 to the value.
#       Add the key and incremented value to 'other_dictionary'.
#    Otherwise, add the key-value pair to 'other_dictionary' as is.
# Note: to loop through key-value pairs in a dictionary, check out the '.items' method.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-234cdbef12134b0d, points: 1, grade: true, solution: false -->
```python
# Tests for Q5

assert other_dictionary

# There are hidden tests that check the contents of 'other_dictionary'.
```

The next question is just about importing the more core data science modules - but we won't start using them yet. 

This is just to test that you have them available and working (this is the only time we'll ask you to write out your own import modules in assignments). 

If you are unable to import them, make sure you have anaconda properly installed (and ask for help in office hours if you can't get them working).

<!-- grade_id: imports, points: 0, grade: false, solution: true -->
```python
# Q6 - Data Science Modules (1 point)

# Import the five core python data science modules, with their typical abbreviations.
#  That is, import {numpy, scipy, pandas, matplotlib.pyplot, sklearn} as {np, sp, pd, plt, skl}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: autograder-tests, points: 1, grade: true, solution: false -->
```python
# Tests for Q6

assert(np)
assert(sp)
assert(pd)
assert(plt)
assert(skl)
```