General instructions!

This is a demonstration notebook of an assignment!

```python
# Imports!
#  In general, we will provide the imports that you need for you.
#  Do not import any other packages that are not explicitly mentioned in the assignment.
#  Make sure you are using Py3.6 with the anaconda distribution - this is the environment we grade in.
import numpy as np
```

<!-- grade_id: cell-63c3a705e90fe682, points: 0, grade: false, solution: true -->
```python
# Q1 - Define a variable, 'a' to have the value of 1.

### BEGIN SOLUTION
a = 1
### END SOLUTION
```

<!-- grade_id: cell-29f67d0cf21aa723, points: 1, grade: true, solution: false -->
```python
# These are the tests for Q1
assert a == 1
```

<!-- grade_id: cell-249c6fc6a3645375, points: 0, grade: false, solution: true -->
```python
# Q2 - Define a variable 'b', that has the value of '2'

### BEGIN SOLUTION
b = '2'
### END SOLUTION
```

<!-- grade_id: cell-b03cc6c2db64df2c, points: 1, grade: true, solution: false -->
```python
# These are the tests for Q2
assert b

### BEGIN HIDDEN TESTS
assert b == '2'
### END HIDDEN TESTS
```

<!-- grade_id: cell-bdbe245776754352, points: 0, grade: false, solution: true -->
```python
# Q3 - Is data science pretty fun? Answer as a boolean called 'ds'

### BEGIN SOLUTION
ds = True
### END SOLUTION
```

<!-- grade_id: cell-645d5b7752dbaf85, points: 1, grade: true, solution: false -->
```python
assert ds
```

<!-- grade_id: cell-fcdbd0de2f88f479, points: 0, grade: false, solution: true -->
```python
# Q4 - Create an array, containing the numbers (as floats) 0-4, inclusively.

### BEGIN SOLUTION
array = np.array(range(5))
### END SOLUTION
```

<!-- grade_id: cell-3ee59982b2db004e, points: 1, grade: true, solution: false -->
```python
assert isinstance(array, np.ndarray)

### BEGIN HIDDEN TESTS
assert np.array_equal(array, np.array([0, 1, 2, 3, 4]))
### END HIDDEN TESTS
```